# Activité de programmation - NSI 1ère

Complétez le document ci-dessous. Vous validerez les cellules au fur-à-mesure en prenant le temps de lire attentivement les consignes.

Certaines cellules contiennent ce texte :
    # YOUR CODE HERE
    raise NotImplementedError()
Vous **remplacerez** ce contenu par votre propre code afin de répondre au problème posé. L'instruction `raise NotImplementedError()` est conçue pour provoquer une erreur si vous ne traitez pas la question.

Une fois le document complété, si le professeur le demande, cliquer sur le bouton **submit** dans l'onglet *assignments*.

Ce *notebook* peut être considéré comme un cahier interactif. N'hésitez pas à vous l'approprier en ajoutant des cellules de texte ou de code selon vos besoins !

# Exercice 1: Les vulnérabilités dans les systèmes d'exploitation et  les applications en 2018 

* Prendre connaissance de la [documentation python sur les fonctions du module csv vues en cours](https://docs.python.org/fr/3/library/csv.html), plus spécifiquement `csv.reader` et `csv.DictReader`. 
* Ecrire une fonction `lire_fichier(fichier)` qui prend en paramètre un nom de fichier et qui retourne les données tabulées du fichier sous forme d'un tableau de tableaux.  
Le fichier de travail est un fichier texte `vulnerabilites.csv` (*FILE1*) qui présente le nombre de vulnérabilités découverts en 2018 dans les systèmes d'exploitations et applications. Les différents champs sont délimités par un point virgule. Voici un exemple d'entrées dans ce fichier:  
Product Name;Vendor Name;Product Type;Total  
Mac Os X;Apple;OS;2208  
Le tableau retourné par votre fonction aura la forme:  
`[['Product Name','Vendor Name','Product Type','Total'], ['Mac Os X','Apple','OS','2208'], ...]`  
La fonction d'affichage `print_table()` peut être éventuellement utilisée (*non obligatoire*)  
* Combien d'enregistrement comporte cette table?

```python
from pathlib import Path


#Constantes
FILE1 = './data/vulnerabilites.csv'
FILE2 = './data/groupe1.csv'

def print_table(table):
    """
    Affiche une table avec un minimum de formattage.
    Entrée: tableau de tableaux
    """
    for ligne in table:
        print(f'{ligne[0]:<28}', end="")
        print(f'{ligne[1]:>15}{ligne[2]:>15}{ligne[3]:>7}')
```

<!-- grade_id: cell-9e80ba27b406a54d, points: 0, grade: false, solution: true -->
```python
import csv


def lire_fichier(fichier):
    """
    Lit un fichier csv depuis le répertoire courant et retourne les données tabulées sous la forme d'un tableau de tableaux
    fichier: nom de fichier (string)
    """
    assert Path(fichier).is_file(), "Erreur: fichier absent"
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c0b71f25090bbf74, points: 4, grade: true, solution: false -->
```python
assert lire_fichier(FILE1)[3][0] == 'Thunderbird', "Erreur: la fonction ne passe pas les tests"
```

```python
print_table(lire_fichier(FILE1))
```

# Exercice 2 Modifier le type d'un attribut
Lors de la création de la structure de données associée à l'ensemble des données tabulées, **tous les champs enregistrés sont de type chaine de caractères** (*string*). On souhaiterait cependant avoir quelques fois avoir des attributs entiers ou flottants.  
* Modifier légèrement la fonction précédente de manière à ce qu'elle retourne une table dont le dernier attribut `Total` soit de type entier.

<!-- grade_id: cell-a96b3e4b9cfd24dd, points: 0, grade: false, solution: true -->
```python
import csv

def lire_fichier2(fichier):
    """
    Lit un fichier csv depuis le répertoire courant et retourne les données tabulées sous la forme d'un tableau de tableaux
    fichier: nom de fichier (string)
    """
    assert Path(fichier).is_file(), "Erreur: fichier absent"
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d99328c34600aa87, points: 4, grade: true, solution: false -->
```python
assert lire_fichier2(FILE1)[3][3] == 872, "Erreur: la fonction ne passe pas les tests"
```

* Quel est le nombre total de vulnérabilités relevé en 2018 ?

```python
table = lire_fichier2(FILE1)
somme = 0
#Votre code
```

# Exercice 3 Ajouter un attribut ou un enregistrement à une table
Dans le répertoire `./data` on dispose d'un fichier `groupe1.csv` (*FILE2*) qui regroupe les données tabulées suivantes:  

| Nom  | LV1 | NSI | Math |
|------|-----|-----|------|
| Alex | 17  | 16  | 16   |
| Zoé  | 15  | 7   | 9    |
| Max  | 13  | 19  | 14   |

1. Ecrire une fonction `lire_table` qui prend en paramètre un nom de fichier (*string*) et qui retourne un tableau de n-uplets nommés (*des dictionnaires ici*) qui représente la table ci-dessus. **Le caractère séparateur est la virgule dans cet exercice**.

<!-- grade_id: cell-3e3f885a901ce64d, points: 0, grade: false, solution: true -->
```python
import csv
from pathlib import Path


#On utilise le canevas du cours

def lire_table(fichier):
    """
    Retourne les données tabulées de 'fichier' sous la forme d'un tableau de dictionnaire
    fichier: nom du fichier csv (string) présent dans le dossier data
    """
    
    assert Path(fichier).is_file(), "Fichier absent"
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6d04caa523bdb81b, points: 4, grade: true, solution: false -->
```python
assert lire_table(FILE2)[0]['Nom'] == 'Alex', "Erreur: la fonction ne passe pas les tests"
```

2. Ecrire l'instruction permettant de modifier la table de données en y ajoutant les notes de Chloé: 17 en math, 19 en LV1 et 18 en NSI.

```python
table = lire_table(FILE2)
#Votre code
```

```python
assert table[3]['Nom'] == 'Chloé', "Erreur: la fonction ne passe pas les tests"
assert table[3]['NSI'] == '18', "Erreur: la fonction ne passe pas les tests"
```

3. On voudrait ajouter une colonne contenant les moyennes de chaque élève, comme dans la table ci-dessous:

| Nom   | LV1 | NSI | Math | Moyenne |
|-------|-----|-----|------|---------|
| Alex  | 17  | 16  | 16   | 16.3    |
| Zoé   | 15  | 7   | 9    | 10.3    |
| Max   | 13  | 19  | 14   | 15.3    |
| Chloé | 19  | 18  | 17   | 18.0    |

Comme nos données tabulées sont représentées par un tableau de n-uplets nommés (*des dictionnaires ici*), il suffit de réaliser une boucle sur tous les éléments (*de type dictionnaire*) du tableau et pour chaque élément:  
*  calculer la moyenne de l'élève;
*  ajouter la paire `'Moyenne':valeur` à l'élément en cours  

Ecrire une fonction `ajoute_moyenne` qui prend en paramètre une table (de type tableau de dictionnaire) et qui retourne **une copie de celle-ci** en y ajoutant une colonne moyenne.  
*Note*  
Pour réaliser une copie d'un objet complexe comme notre table, on peut utiliser la fonction `deepcopy` du module `copy`.  
```python
from copy import deepcopy
#....
new_table = deepcopy(table)
```

<!-- grade_id: cell-c9fb16263b9ed19d, points: 0, grade: false, solution: true -->
```python
from copy import deepcopy


def ajoute_moyenne(t):
    """
    Ajoute un champ 'Moyenne' aux données tabulées et retourne un nouveau tableau.
    t: tableau de dictionnaires
    """
    new_t = deepcopy(t)
    # YOUR CODE HERE
    raise NotImplementedError()
    return new_t
```

<!-- grade_id: cell-890e2482c16b6678, points: 0, grade: true, solution: false -->
```python
assert ajoute_moyenne(table)[0]['Moyenne'] == '16.3', "Erreur: la fonction ne passe pas les tests"
assert ajoute_moyenne(table)[1]['Moyenne'] == '10.3', "Erreur: la fonction ne passe pas les tests"
```

4. (*A faire à la maison*) On voudrait rajouter un enregistrement qui donne la moyenne par matière, comme dans la table suivante:  

| Nom     | LV1  | NSI  | Math | Moyenne |
|---------|------|------|------|---------|
| Alex    | 17   | 16   | 16   | 16.3    |
| Zoé     | 15   | 7    | 9    | 10.3    |
| Max     | 13   | 19   | 14   | 15.3    |
| Chloé   | 19   | 18   | 17   | 18      |
| moy_mat | 16.0 | 15.0 | 14.0   | 15.0    |

Ecrire une fonction `ajoute_moy_matieres` qui prend en paramètre une table (de type tableau de dictionnaire) et qui retourne **une copie de celle-ci** en y ajoutant une ligne comportant les moyennes par matière. 

<!-- grade_id: cell-b891da3664d6825e, points: 0, grade: false, solution: true -->
```python
from copy import deepcopy


def ajoute_moy_matieres(t):
    """
    Ajoute un nouvel enregistrement (moyenne par matière) aux données tabulées et retourne un nouveau tableau.
    t: tableau de dictionnaires
    """
    #Création d'un nouveau tableau, copie de t, avec 1 élément en plus (qui contiendra la moyenne/matière)
    new_t = deepcopy(t) + [{}]
    #Création d'un dictionnaire (dernière ligne de la table) qui contiendra la paire matiere:moyenne;
    #Le champ 'Nom' aura pour valeur 'moy_mat' (voir tableau de l'énoncé)
    moy_mat = {'Nom':'moy_mat'}
    
    # YOUR CODE HERE
    raise NotImplementedError()
    return new_t
```

<!-- grade_id: cell-928f410b120fa7b1, points: 4, grade: true, solution: false -->
```python
table = lire_table(FILE2)
table.append({'Nom':'Chloé', 'Math':'17', 'LV1':'19', 'NSI':'18'})
table = ajoute_moyenne(table)

assert ajoute_moy_matieres(table)[-1]['NSI'] == '15.0', "Erreur: la fonction ne passe pas les tests"
```