# Activité de programmation - NSI 1ère

Complétez le document ci-dessous. Vous validerez les cellules au fur-à-mesure en prenant le temps de lire attentivement les consignes.

Certaines cellules contiennent ce texte :
    # YOUR CODE HERE
    raise NotImplementedError()
Vous **remplacerez** ce contenu par votre propre code afin de répondre au problème posé. L'instruction `raise NotImplementedError()` est conçue pour provoquer une erreur si vous ne traitez pas la question.

Une fois le document complété, si le professeur le demande, cliquer sur le bouton **submit** dans l'onglet *assignments*.

Ce *notebook* peut être considéré comme un cahier interactif. N'hésitez pas à vous l'approprier en ajoutant des cellules de texte ou de code selon vos besoins !

# Exercice 1
Soit le tableau `t = [[1, 2, 3], [4, 5, 6], [7, 8, 9]]`.
On considère les instructions suivantes:
```python
a = t[1]
b = a[2]
```
Que valent `a` et `b`?

# Exercice 2
Ecrire une fonction f qui prend un tableau de tableaux `t` en paramètre. La fonction effectuera une copie `t2`de `t` puis remplacera le dernier élément de chacun des tableaux de `t2` par la moyenne des éléments précédents et enfin retournera `t2`. Par exemple, pour le tableau de l'exercice 1, on aurait:
```python
print(f(t))
[[1, 2, 1.5], [4, 5, 4.5], [7, 8, 7.5]]
```
Attention, la tableau `t` passé en paramètre a un nombre quelconque de colonne et que chaque ligne a la même taille!

<!-- grade_id: cell-1a1c34b68a943c77, points: 0, grade: false, solution: true -->
```python
from copy import deepcopy

def f(t):
    """
    """
    
    t2 = deepcopy(t)
    #Pour une copie de tableau de tableaux, utiliser une fonction dédiée
    # YOUR CODE HERE
    raise NotImplementedError()
    return t2
```

<!-- grade_id: cell-68926d5c86da5478, points: 4, grade: true, solution: false -->
```python
assert f([[1, 2, 3], [4, 5, 6], [7, 8, 9]]) == [[1, 2, 1.5], [4, 5, 4.5], [7, 8, 7.5]]
```